**Student name and BC username:**

**Other contributors (up to 2) and BC username**

# <font face = "times"><center>  EC1151 - Homework 1 (Due Sep 9th, 3PM) 

<font face="times">
# Instructions

+ You must fill out this Jupyter notebook and return it by e-mail to baisihad@bc.edu

+ Double click on the cells to edit them

+ For questions that need mathematical notation, try to use $\LaTeX$ if you can.

+ Sometimes, below a question that requires an answer there will be a cell with several `assert` statements. If you can run that cell without generating error, your answer has been validated and you are guaranteed the points.

+ This assignment will be graded out of 110 points, because of the extra 10 points for $\LaTeX$ at the bottom. However, your maximum recorded grade will be capped at 100.

+ If you're stuck, try posting a question in the Piazza forum. Active contributors will be rewarded.

+ Finally, this is <b>experimental</b>. I'd love to hear what you think about this homework format. You can e-mail me with your thoughts, or (preferrably) post in the Piazza forum (possibly anonymously).

# <font face = "times">Theory 

## <font face = "Times"> Q1: Sampling (20 points)</font>

Remember to explain your reasoning, even for the question of mathematical content.

<font face = "Times">**Q1A)** In how many ways can we *sample from* 10 entrees and 10 deserts? (That is, how many different meals can we make?) 

<!-- grade_id: q1_a_sol, points: 4, grade: true, solution: true -->
YOUR ANSWER HERE

<font face = "Times">**Q1B)** If we pick two meals out of three entrees, we can make three different courses. For example, out of {chicken, beef, fish} we can choose {chicken, beef}, {chicken, fish}, {beef, fish}.

If we pick five meals out of ten entrees, how many different courses can we make?

Hint: How many choices do you have for the first meal? How many choices are left for the second meal?

<!-- grade_id: q1_b_sol, points: 4, grade: true, solution: true -->
YOUR ANSWER HERE

<font face = "times">**Q1C)** A `bit` is a binary digit: it can be 0 or 1. A `byte` is a collection of `8` bits. For example, 10101010 is a `byte`, and so is 00011100. How many distinct bytes can we create out of 8 `bits`?

Hint: A `nibble` is a collection of `4` bytes, like 1011 or 1111 or 0001. There are `16` different possible nibbles.

<!-- grade_id: q1_c_sol, points: 6, grade: true, solution: true -->
YOUR ANSWER HERE

<font face = "times">**Q1D)** A *palindrome* is a sequence of characters that reads the same way forwards and backwards. Boring examples include "race car" and "taco cat". For a less boring example see Demetri Martin's "[Dammit I'm mad](http://fuckyeahdemetrimartin.tumblr.com/post/1382923819/simonfm-dammit-im-mad-a-poem-that-is-a)" poem. 

Is a word with 2 letters more likely to be a palindrome than a word with 3 letters in it?


<!-- grade_id: q1_d_sol, points: 6, grade: true, solution: true -->
YOUR ANSWER HERE

<font face = "times">**Q1E)** [No points - Extra] This is the famous [Boy of Girl Paradox](https://en.wikipedia.org/wiki/Boy_or_Girl_paradox).

+ Mr. Jones has two children. The older child is a girl. What is the probability that both children are girls?
+ Mr. Smith has two children. At least one of them is a boy. What is the probability that both children are boys?

Try thinking what answers you'd give before reading the article. 

```python
""" YOUR ANSWER HERE """
```

## Q2 - Coin flipping experiment (30 points)

<font face = "times"><b>Q2A)</b> If you toss a coin three times, what is the probability of getting 2 heads and 1 tail? Explain your reasoning.

<!-- grade_id: q2_a_sol, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

<font face = "times">**Q2B)** For the next experiment, we'll repeatedly toss three coins. We'll call a *success* if the event of getting "exactly two heads and one tail" is true. For example, `HTH` is a *sucess*, and so is `HHT`; however `TTT` is a failure.

Repeat the following experiment **10** times.

+ Get three coins and toss them. Record if you got a *sucess* or not.

Compute the ratio of *successes* over total experiments. (i.e., if you had 3 *successes*, you should answer $\frac{3}{10} = 0.3$).

""" Double-click on this text, erase it and write your own answer """ 

<font face = "times">
Perform the previous experiment an additional 15 times. Compute the ratio of *successes* over total experiments. What is the ratio of successes over total repetitions (now 25)?

""" Double-click on this text, erase it and write your own answer """ 

<font face = "times">

Perform the previous experiment an additional 25 times (or more!).
Compute the ratio of *successes* over total experiments. What is the ratio of successes over total repetitions (now at least 50)?

""" Double-click on this text, erase it and write your own answer """ 

<font face = "times">**Q2C)** What do you observe about the numbers you found in question b? Try to relate them to your answer in question a.

In particular, if you continued repeating your experiment trillions of times, what is the approximate ratio of "successes / total repetitions" that you would have gotten?

<!-- grade_id: q2_b_sol, points: 20, grade: true, solution: true -->
YOUR ANSWER HERE

#  Programming Practice

## Q1 - Printing (5 points)

Use the `print` function to print out the message "Hello World!"

<!-- grade_id: print1, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Q2 - Functions (20 points)

Define a function `two_power(x)` that takes a number `x` and outputs $2^x$.

<!-- grade_id: two_to_the_power_of, points: 10, grade: false, solution: true -->
```python
def two_power(x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
two_power  = lambda x: 2*x
```

<!-- grade_id: two_to_the_power_of_tests, points: 10, grade: true, solution: false -->
```python
from numpy.random import randint
assert two_power(10) == 20
assert two_power(0) == 0
assert two_power(r) == 2*r
```

Define a function `reciprocal` that takes the number `x` and outputs $1/x$. (I promise not to use it for x = 0)

<!-- grade_id: reciprocal, points: 10, grade: false, solution: true -->
```python
def reciprocal(x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: reciprocal_tests, points: 10, grade: true, solution: false -->
```python
r = randint(10, 100)
assert reciprocal(1) == 1
assert reciprocal(10) == 0.1
assert reciprocal(r) == 1/r
```

## Q2 - Numpy (25 points)

Create a variable `x` containing 100 randomly-generated numbers, uniformly distributed between 0 and 1.

<!-- grade_id: random, points: 25, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

[Extra - Zero points] Use the `numpy` function `mean` to take the mean of the 100 numbers, and store that in the variable `x_mean`.

<!-- grade_id: mean, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# Extra (10 points)

## Writing in $\LaTeX$

Reproduce the following using $\LaTeX$ code

$$a) \ \ x^{2} \qquad 
b) \ \ \frac{1}{2} \qquad 
c) \ \ \Omega \qquad
d) \ \ \Gamma $$

<!-- grade_id: l1, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: l2, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: l3, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: l4, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: l5, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Remember to enclose $\LaTeX$ code in single or double dollar signs, as in \$\\LaTeX\$.