**Student name and BC username**

NAME HERE

**Collaborators names and BC usernames (max 2 collaborators)**

+ NAME 1 

+ NAME 2

---

# <font face = "times"><center>  EC1151 - Homework 1 (Due Sep 9th, 3PM) 

<font face="times">
# Instructions

+ You must fill out this Jupyter notebook and return it by e-mail to baisihad@bc.edu

+ Double click on the cells to edit them

+ For questions that need mathematical notation, try to use $\LaTeX$ if you can.

+ Sometimes, below a question that requires an answer there will be a cell with several `assert` statements. If you can run that cell without generating error, your answer has been validated and you are guaranteed the points.

+ This assignment will be graded out of 110 points, because of the extra 10 points for $\LaTeX$ at the bottom. However, your maximum recorded grade will be capped at 100.

+ If you're stuck, try posting a question in the Piazza forum. Active contributors will be rewarded.

+ Finally, this is <b>experimental</b>. I'd love to hear what you think about this homework format. You can e-mail me with your thoughts, or (preferrably) post in the Piazza forum (possibly anonymously).

---

# <font face = "times">Theory 

## <font face = "Times"> Q1: Sampling (35 points)</font>

Remember to explain your reasoning, even for the question of mathematical content.

<font face = "Times">**Q1A)** In how many ways can we *sample from* 10 entrees and 10 deserts? (That is, how many different meals can we make?) 

<!-- grade_id: q1a, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

<font face = "Times">**Q1B)** If we pick two meals out of three entrees, we can make three different courses. For example, out of {chicken, beef, fish} we can choose {chicken, beef}, {chicken, fish}, {beef, fish}.

If we pick five meals out of ten entrees, how many different courses can we make?

Hint: How many choices do you have for the first meal? How many choices are left for the second meal?

<!-- grade_id: q1b, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

<font face = "times">**Q1C)** A `bit` is a binary digit: it can be 0 or 1. A `byte` is a collection of `8` bits. For example, 10101010 is a `byte`, and so is 00011100. How many distinct bytes can we create out of 8 `bits`?

Hint: A `nibble` is a collection of `4` bytes, like 1011 or 1111 or 0001. There are `16` different possible nibbles.

<!-- grade_id: q1c, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

<font face = "times">**Q1D)** A *palindrome* is a sequence of characters that reads the same way forwards and backwards. Boring examples include "race car" and "taco cat". For a less boring example see Demetri Martin's "[Dammit I'm mad](http://fuckyeahdemetrimartin.tumblr.com/post/1382923819/simonfm-dammit-im-mad-a-poem-that-is-a)" poem. 

Is a word with 2 letters more likely to be a palindrome than a word with 3 letters in it?

<!-- grade_id: q1d, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

<font face = "times">**Q1E)** [No points - Extra] This is the famous [Boy of Girl Paradox](https://en.wikipedia.org/wiki/Boy_or_Girl_paradox).

+ Mr. Jones has two children. The older child is a girl. What is the probability that both children are girls?
+ Mr. Smith has two children. At least one of them is a boy. What is the probability that both children are boys?

Try thinking what answers you'd give before reading the article. 

YOUR ANSWER HERE

## Q2 - Coin flipping experiment (30 points)

<font face = "times"><b>Q2A)</b> If you toss a coin three times, what is the probability of getting 2 heads and 1 tail? Explain your reasoning.

<!-- grade_id: q2a, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

<font face = "times">**Q2B)** For the next experiment, we'll repeatedly toss three coins. We'll call a *success* if the event of getting "exactly two heads and one tail" is true. For example, `HTH` is a *sucess*, and so is `HHT`; however `TTT` is a failure.

Repeat the following experiment **10** times.

+ Get three coins and toss them. Record if you got a *sucess* or not.

Compute the ratio of *successes* over total experiments. (i.e., if you had 3 *successes*, you should answer $\frac{3}{10} = 0.3$).

<!-- grade_id: q2b1, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

<font face = "times">
Perform the previous experiment an additional 15 times. Compute the ratio of *successes* over total experiments. What is the ratio of successes over total repetitions (now 25)?

<!-- grade_id: q2b2, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

<font face = "times">

Perform the previous experiment an additional 25 times (or more!).
Compute the ratio of *successes* over total experiments. What is the ratio of successes over total repetitions (now at least 50)?

<!-- grade_id: q2b3, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

<font face = "times">**Q2C)** What do you observe about the numbers you found in part b? Try to relate them to your answer in question a.

In particular, if you continued repeating your experiment trillions of times, what is the approximate ratio of "successes / total repetitions" that you would have gotten?

<!-- grade_id: q2c, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

---

<font face="times">
#  Programming Practice

<font face="times">
## Printing (5 points)

<font face="times">
**Q3)** Use the `print` function to print out the message "Hello World!"

<!-- grade_id: q3, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<font face="times">
## Functions (20 points)

<font face="times">
**Q4)** Define a function `triple(x)` that takes `x` and outputs three times `x`.

<!-- grade_id: q4, points: 0, grade: false, solution: true -->
```python
def triple(x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: q4t, points: 5, grade: true, solution: false -->
```python
assert two_power(0) == 1
assert two_power(1) == 2
assert two_power(4) == 16
assert two_power(10) == 1024
```

<font face="times">
**Q5)** Define a function `two_power(x)` that takes a number `x` and outputs $2^x$ (i.e., two to the power of `x`)

<!-- grade_id: q5, points: 5, grade: false, solution: true -->
```python
def two_power(x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: q5t, points: 5, grade: true, solution: false -->
```python
assert two_power(0) == 1
assert two_power(1) == 2
assert two_power(4) == 16
assert two_power(10) == 1024
```

<font face="times">
**Q6)** Define a function `summation(x, y)` that takes two arguments: `x` and `y`, and then outputs their sum.

<!-- grade_id: q6, points: 0, grade: false, solution: true -->
```python
def summation(x, y):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: q6t, points: 5, grade: true, solution: false -->
```python
assert summation(1, 1) == 2
assert summation(1, 2) == 3
assert summation(3, 1.5) == 4.5
assert summation(60, 60) == 120
```

<font face="times">
**Q7)** Define a function `constant_five()` that takes *no* arguments and always outputs the number $5$.

<!-- grade_id: q7, points: 0, grade: false, solution: true -->
```python
def constant_five():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: q7t, points: 5, grade: true, solution: false -->
```python
assert constant_five() == 5
```

<font face="times">
## Numpy (10 points)

<font face="times">
**Q8)** Suppose you want to compute the sine, cosine and tangent of a number. You know that the `numpy` library has functions called `cos`, and `sin` and `tan` that would fit your needs. How can you import those functions *in a way that you wouldn't need to write `numpy.` as a prefix*?

<!-- grade_id: q8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: q8t, points: 10, grade: true, solution: false -->
```python
# Note that we're not using the prefix numpy. here!
assert cos(0) == 1
assert sin(0) == 0
assert tan(0) == 0
```

# $\LaTeX$ (10 points)

<font face="times">
**Q9)** Reproduce the following using $\LaTeX$ code

<img src = "hw1_latex.png", width = 500>

Part a (exponent)

<!-- grade_id: q9a, points: 2.5, grade: true, solution: true -->
YOUR ANSWER HERE

Part b (fraction)

<!-- grade_id: q9b, points: 2.5, grade: true, solution: true -->
YOUR ANSWER HERE

Part c (Greek letter capital Omega)

<!-- grade_id: q9c, points: 2.5, grade: true, solution: true -->
YOUR ANSWER HERE

Part d (Greek letter capital Gamma)

<!-- grade_id: q9d, points: 2.5, grade: true, solution: true -->
YOUR ANSWER HERE